<?xml version="1.0"?>
<exercises xml:id="conditionals_write-code-questions">
  <title>Write Code Questions</title>

        <exercise label="cndtnl-wc-fivea">
          <statement>
            <p>Fix the errors in the code, and change it to use only <em>one</em> <c>if</c> statement. The code should print
                            <q>The number is 5</q> when the number is 5, and should print <q>The number is NOT 5</q> when it is not.</p>
          </statement>
          <program interactive="activecode" language="python" xml:id="cndtnl-wc-fiveq_editor">
            <input>
x = 5
if x == 5:
print("The number is 5")
if x != 5:
print("The number is NOT 5")
        </input>
          </program>
          <solution>
    <program xml:id="cndtnl-wc-fivea_editor" language="python"><input>
# Initialize x and set condition for when x = 5
# Use "else" to take care of condition when x is not 5.
x = 5
if x == 5:
    print("The number is 5") # indent the body of if
else:
    print("The number is NOT 5") # indent the body of else
        </input></program>
                    </solution>
        </exercise>

        <exercise label="cndtnl-wc-hiq">
          <statement>
            <p>Complete the conditional below so the words <q>Not three</q> are printed if <c>x</c> does not
                    equal 3 and <q>It is three</q> is printed otherwise.</p>
          </statement>
          <program xml:id="cndtnl-wc-hiq_editor" interactive="activecode" language="python">
            <input>
x = 4
if x
        </input>
          </program>
        </exercise>

        <exercise label="cndtnl-wc-payq">
          <statement>
            <p>Complete the pay computation as follows: When the employee works less than or equal to 40 hours a week, pay them $10 per hour. Otherwise, pay them $10 and hour for the first 40 hours and 1.5 times the hourly rate for hours worked above 40 hours. Then set <c>grossPay</c> equal to the  amount an employee would be paid for working 45 hours.</p>
          </statement>
          <program interactive="activecode" language="python" xml:id="cndtnl-wc-payq_editor">
            <input>
hours = 45
rate = 10
overtimeRate =  # your code here
grossPay = 0
if hours  # your code here
    grossPay = hours * rate
else:
    overTime =  # your code here
    grossPay =  # your code here
print(grossPay)

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(grossPay, 475.0)

myTests().main()
        </input>
          </program>
          <solution>
    
    <program xml:id="cndtnl-wc-paya_editor" language="python"><input>
# Initializing variables
hours = 45
rate = 10
# overtimeRate is 1.5 the rate amount
overtimeRate = rate * 1.5
grossPay = 0
# Begin conditional to see if hours are within regular pay
if hours &lt;= 40:  # for 40 hours or fewer...
    grossPay = hours * rate
else:
    overTime = hours - 40 # calculate hours over 40
    # Pay will equal the regular rate for 40 hours,
    # plus the overtime rate for the extra hours
    grossPay = (rate * 40) + (overTime * overtimeRate)
# Print the final pay
print(grossPay)
        </input></program>
                    </solution>
        </exercise>

        
        <exercise label="cndtnl-wc-distanceq">
          <statement>
            <p>Write the code to calculate and print the cost of a 14 mile cab ride. If the distance traveled is less than or equal to 5 miles the cost is $2.00 a mile. If the distance is greater than 5 but less than or equal to 12 miles, the cost is $1.50, and if the distance traveled is more than 12 miles the cost is $1.25 a mile. Assign the final cost to the variable <c>total</c>. (Your code should work for any mileage, not just for 14 miles.)</p>
          </statement>
          <program interactive="activecode" language="python" xml:id="cndtnl-wc-distanceq_editor">
            <input>
distance = 14

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(total, 17.5)

myTests().main()
        </input>
          </program>
          <solution>
    <program xml:id="cndtnl-wc-distancea_editor" language="python"><input>
distance = 14 # must initialize variable

if distance &lt;= 5:
    rate = 2.00
elif distance &lt;= 12:
    rate = 1.50
else:
    rate = 1.25

total = distance * rate

print("Total cost of trip:", total)
        </input></program>
                    </solution>
        </exercise>

</exercises>
