<?xml version="1.0"?>
<exercises xml:id="strings_exercises">
  <title>Exercises</title>
  
        <exercise label="str_string_eval">
          <statement>
            <p>What do these evaluate to?</p>
            <p>
              <ol marker="a">
                <li><c>'Python'[1]</c></li>
                <li><c>'Strings are sequences of characters.'[5]</c></li>
                <li><c>len('wonderful')</c></li>
                <li><c>"Mystery"[:4]</c></li>
                <li><c>'p' in 'Pineapple'</c></li>
                <li><c>'apple' in 'Pineapple'</c></li>
                <li><c>'pear' not in 'Pineapple'</c></li>
                <li><c>'apple' &gt; 'pineapple'</c></li>
                <li><c>'pineapple' &lt; 'Peach'</c></li>
              </ol>
            </p>
          </statement>
          
          <solution>
            <p>
              <ol marker="a">
                <li><c>'Python'[1]</c> &#x2192; <c>'P'</c></li>
                <li><c>'Strings are sequences of characters.'[5]</c> &#x2192; <c>'g'</c></li>
                <li><c>len('wonderful')</c> &#x2192; 9</li>
                <li><c>"Mystery"[:4]</c> &#x2192; <c>'Myst'</c></li>
                <li><c>'p' in 'Pineapple'</c> &#x2192; <c>True</c></li>
                <li><c>'apple' in 'Pineapple'</c> &#x2192; <c>True</c></li>
                <li><c>'pear' not in 'Pineapple'</c> &#x2192; <c>True</c></li>
                <li><c>'apple' &gt; 'pineapple'</c> &#x2192; <c>False</c></li>
                <li><c>'pineapple' &lt; 'Peach'</c> &#x2192; <c>False</c></li>
              </ol>
            </p>
          </solution>
        </exercise>

        <exercise label="str_modify_strings">
          <statement>
            <p>In Robert McCloskey's
                    book <em>Make Way for Ducklings</em>, the names of the ducklings are Jack, Kack, Lack,
                    Mack, Nack, Ouack, Pack, and Quack.  This loop tries to output these names in order.</p>
            <program language="python">
              <input>
prefixes = "JKLMNOPQ"
suffix = "ack"

for p in prefixes:
    print(p + suffix)
</input>
            </program>
            <p>Of course, that's not quite right because Ouack and Quack are misspelled. Can you fix it?</p>
          </statement>
          <program xml:id="str_modify_strings_editor" interactive="activecode" language="python">
            <input>

prefixes = "JKLMNOPQ"
suffix = "ack"

for p in prefixes:
    print(p + suffix)

# Fix the loop to get the correct output.
</input>
<tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):

        o = self.getOutput()
        code = self.getEditorText()
        self.assertIn("if", code, "Needs a conditional.")
        self.assertIn("for", code, "Needs a loop.")
        self.assertIn("Jack", o, "J + ack = Jack")
        self.assertIn("Kack", o, "K + ack = Kack")
        self.assertIn("Lack", o, "L + ack = Lack")
        self.assertIn("Mack", o, "M + ack = Mack")
        self.assertIn("Nack", o, "N + ack = Nack")
        self.assertIn("Ouack", o, "Don't forget the misspellings. Quack is required.")
        self.assertIn("Pack", o, "P + ack = Pack")
        self.assertIn("Quack", o, "Don't forget the misspellings. Quack is required.")
        self.assertNotIn("Oack", o, "Account for the misspellings. Qack should not be in output.")
        self.assertNotIn("Qack", o, "Account for the misspellings. Qack should not be in output.")
myTests().main()
        </tests>
          </program>
        </exercise>

        
        <exercise label="str_count_letter_e">
          <statement>
            <p>Assign to a variable in your program a triple-quoted string that contains your favorite paragraph of text - perhaps a poem, a speech, instructions to bake a cake, some inspirational verses, etc.</p>
            <p>Write a function that counts the number of alphabetic characters in your text and then keeps track of how many are the letter 'e'.  Your function should print an analysis of the text like this:</p>
            <pre>Your text contains 243 alphabetic characters, of which 109 (44.8%) are 'e'.</pre>
          </statement>
          <program interactive="activecode" language="python" xml:id="str_count_letter_e_editor">
            <input>
def count(paragraph):
    """Count alpha characters in paragraph.
    
    This function also counts the number of times
    the letter "e" appears in the paragraph, and
    prints the percentage of characters that are "e".
    """
    # your code here
  </input>
  <tests>

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        string1 = "e"
        string2 = "eieio"
        string3 = "eeeeeeeeeeee"
        string4 = "elephant"
        self.assertEqual(count(string1), 1, "Just one e")
        self.assertEqual(count(""), 0, "Empty string")
        self.assertEqual(count(string2), 2, "Two")
        self.assertEqual(count(string3), 12, "Twelve")
        self.assertNotEqual(count(string4), 3, "Has two Es")
myTests().main()
        </tests>
          </program>
          <solution>
            <program xml:id="str_q3_answer" language="python">
              <input>
def count(paragraph):
    number_of_e = 0
    total_chars = 0
    for a_char in paragraph:
        if a_char.isalpha():
            total_chars = total_chars + 1
            if a_char == 'e':
                number_of_e = number_of_e + 1

    percent_with_e = (number_of_e / total_chars)
    print(f"Your text contains {total_chars} alphabetic characters of which {number_of_e} ({percent_with_e:.2%}) are 'e'.")



para = '''
"If the automobile had followed the same development cycle as the computer, a
Rolls-Royce would today cost $100, get a million miles per gallon, and explode
once a year, killing everyone inside."
-Robert Cringely
'''

count(para)
        </input>
            </program>
          </solution>
        </exercise>

        <!--
        <exercise label="str_8_4">
          <statement>
            <p>Print out a neatly formatted multiplication table, up to 12 x 12.</p>
          </statement>
          <program xml:id="str_8_4_editor" interactive="activecode" language="python">
            <input>

        </input>
          </program>
        </exercise>
        -->

        <exercise label="str_num_digits">
          <statement>
            <p>Write a function that will return the number of digits in an integer.</p>
          </statement>
          <program interactive="activecode" language="python" xml:id="str_num_digits_editor">
            <input>
def numDigits(n):
    """Return number of digits in given integer n."""
    # your code here
  </input>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(numDigits(2),1,"Tested numDigits on input of 2")
      self.assertEqual(numDigits(55),2,"Tested numDigits on input of 55")
      self.assertEqual(numDigits(1352),4,"Tested numDigits on input of 1352")
      self.assertEqual(numDigits(444),3,"Tested numDigits on input of 444")

myTests().main()
        </tests>
          </program>
          <solution>
            <program xml:id="str_num_digits_answer" language="python">
              <input>
def numDigits(n):
    n_str = str(n)
    return len(n_str)


print(numDigits(50))
print(numDigits(20000))
print(numDigits(1))
        </input>
            </program>
          </solution>
        </exercise>

        <exercise label="str_reverse">
          <statement>
            <p>Write a function that reverses its string argument.</p>
          </statement>
          <program xml:id="str_reverse_editor" interactive="activecode" language="python">
            <input>
def reverse(astring):
    """Return reverse of the given string argument."""
    # your code here
  </input>
  <tests>

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(reverse("happy"),"yppah","Tested reverse on input of 'happy'")
      self.assertEqual(reverse("Python"),"nohtyP","Tested reverse on input of 'Python'")
      self.assertEqual(reverse(""),"","Tested reverse on input of ''")




myTests().main()
        </tests>
          </program>
        </exercise>

        <exercise label="str_remove_letter">
          <statement>
            <p>Write a function that removes all occurrences of a given letter from a string.</p>
          </statement>
          <program xml:id="str_remove_letter_editor" interactive="activecode" language="python">
            <input>
def remove_letter(the_letter, the_string):
    """Remove all occurrences of the_letter from the_string."""
    # your code here
  </input>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_letter("a","apple"),"pple","Tested remove_letter on inputs of 'a' and 'apple'")
      self.assertEqual(remove_letter("a","banana"),"bnn","Tested remove_letter on inputs of 'a' and 'banana'")
      self.assertEqual(remove_letter("z","banana"),"banana","Tested remove_letter on inputs of 'z' and 'banana'")


myTests().main()
        </tests>
          </program>
        </exercise>

        <exercise label="str_palindrome">
          <statement>
            <p>A palindrome is a word that is spelled the same backwards and forwards. For example, <q>racecar</q> and <q>madam</q> are palindromes. Write a function that recognizes palindromes. (Hint: use your <c>reverse</c> function to make this easy!).</p>
          </statement>
          <program interactive="activecode" language="python" xml:id="str_palindrome_editor">
            <input>
def is_palindrome(my_str):
    """Returns True if my_str is same forward as backward."""
    # your code here
  </input>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(is_palindrome("abba"),True,"Tested is_palindrome on input of 'abba'")
        self.assertEqual(is_palindrome("abab"),False,"Tested is_palindrome on input of 'abab'")
        self.assertEqual(is_palindrome("straw warts"),True,"Tested is_palindrome on input of 'straw warts'")
        self.assertEqual(is_palindrome("a"),True,"Tested is_palindrome on input of 'a'")
        self.assertEqual(is_palindrome(""),True,"Tested is_palindrome on input of ''")




myTests().main()
        </tests>
          </program>
          <solution>
            <program xml:id="str_palindrome_answer" language="python">
              <input>
def reverse(my_str):
    reversed = ''
    for char in my_str:
        reversed = char + reversed
    return reversed

def is_palindrome(my_str):
    """Returns True if my_str is same forward as backward."""
    if my_str == reverse(my_str):
        return True
    else:
        return False

print(is_palindrome('racecar')) # True
print(is_palindrome('pacecar')) # False
print(is_palindrome('straw warts')) # True
print(is_palindrome('a')) # True
print(is_palindrome('')) # True
        </input>
            </program>
          </solution>
        </exercise>

        <!--
        <exercise>
          <statement>
        <p>Write a function that counts how many non-overlapping occurences of a substring appear in a string.</p>
      </statement>
      <program xml:id="str_8_9" interactive="activecode" language="python">
          <input>
def count(substr,theStr):
    # your code here
        </input>
        <tests>


from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

      def testOne(self):
          self.assertEqual(count("is","Mississippi"),2,"Tested count on inputs of 'is' and 'Mississippi'")
          self.assertEqual(count("an","banana"),2,"Tested count on inputs of 'an' and 'banana'")
          self.assertEqual(count("ana","banana"),1,"Tested count on inputs of 'ana' and 'banana'")
          self.assertEqual(count("nana","banana"),1,"Tested count on inputs of 'nana' and 'banana'")
          self.assertEqual(count("nanan","banana"),0,"Tested count on inputs of 'nanan' and 'banana'")
          self.assertEqual(count("aaa","aaaaaa"),2,"Tested count on input of 'aaa' and 'aaaaaa'")




myTests().main()
        </tests>
        </program>

    </exercise>
        -->

        <exercise label="str_remove_first_substr">
          <statement>
            <p>Write a function that removes the first occurrence of a string from another string. For example, <c>remove("on", "London")</c> would return <c>"Ldon"</c>. Hint: use slices.</p>
          </statement>
          <program interactive="activecode" language="python" xml:id="str_remove_first_substr_editor">
            <input>
def remove(substr, the_str):
    """Remove first occurrence of substr from the_str."""
    # your code here
  </input>

  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

        def testOne(self):
            self.assertEqual(remove("an","banana"),"bana","Tested remove on inputs of 'an' and 'banana'")
            self.assertEqual(remove("cyc","bicycle"),"bile","Tested remove on inputs of 'cyc' and 'bicycle'")
            self.assertEqual(remove("iss","Mississippi"),"Missippi","Tested remove on inputs of 'iss' and 'Mississippi'")
            self.assertEqual(remove("egg","bicycle"),"bicycle","Tested remove on inputs of 'egg' and 'bicycle'")

myTests().main()
        </tests>
          </program>
          <solution>
            <program xml:id="str_remove_first_substr_answer" language="python">
              <input>
def remove(substr,theStr):
    index = theStr.find(substr)
    if index &lt; 0: # substr doesn't exist in theStr
        return theStr
    return_str = theStr[:index] + theStr[index+len(substr):]
    return return_str

# These should all print True
print(remove('an', 'banana') == 'bana')
print(remove('cyc', 'bicycle') == 'bile')
print(remove('iss', 'Mississippi') == 'Missippi')
print(remove('egg', 'bicycle') == 'bicycle')
        </input>
            </program>
          </solution>
        </exercise>

        <exercise label="str_remove_all_substr">
          <statement>
            <p>Write a function that removes all occurrences of a string from another string. For example, <c>remove_all("on", "London")</c> would return <c>"Ld"</c></p>
          </statement>
          <program xml:id="str_remove_all_substr_editor" interactive="activecode" language="python">
            <input>
def remove_all(substr, the_str):
    """Remove all occurrences of substr from the_str."""
    # your code here
  </input>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_all("an","banana"),"ba","Tested remove_all on inputs of 'an' and 'banana'")
      self.assertEqual(remove_all("cyc","bicycle"),"bile","Tested remove_all on inputs of 'cyc' and 'bicycle'")
      self.assertEqual(remove_all("iss","Mississippi"),"Mippi","Tested remove_all on inputs of 'iss' and 'Mississippi'")
      self.assertEqual(remove_all("eggs","bicycle"),"bicycle","Tested remove_all on inputs of 'eggs' and 'bicycle'")

myTests().main()
        </tests>
          </program>
        </exercise>

        <exercise label="str_cipher_encrypt">
          <statement>
            <p>Write a function that implements a substitution cipher.  In a substitution cipher one letter is substituted for another to garble the message.  For example A -&gt; Q, B -&gt; T, C -&gt; G etc. Your function should take two parameters: the message you want to encrypt, and a string that represents the mapping of the 26 letters in the alphabet. This mixed-up alphabet is called the <em>cipher</em>.  Your function should return a string that is the encrypted version of the message. If a character isn't in the alphabet, it should remain as is in the encrypted version.</p>
          </statement>
          <program xml:id="str_cipher_encrypt_editor" interactive="activecode" language="python">
            <input>
def encrypt(message, cipher):
    """Encrypt message using cipher alphabet."""
    # your code here
  
cipher = "badcfehgjilknmporqtsvuxwzy"
encrypted = encrypt("python strings!", cipher)
print(encrypted) # should result in "ozsgpm tsqjmht!"
        </input>
          </program>
        </exercise>

        <exercise label="str_cipher_decrypt">
          <statement>
            <p>
              Write a function that decrypts the message from the previous exercise.  It should also take two parameters.  The encrypted message, and the mixed up alphabet.  The function should return a string that is the same as the original unencrypted message.
            </p>
          </statement>
          <program interactive="activecode" language="python" xml:id="str_cipher_decrypt_editor">
            <input>

        </input>
          </program>
          <solution>
            <program xml:id="str_cipher_decrypt_answer" language="python">
              <input>
def encrypt(message, cipher):
    alphabet = "abcdefghijklmnopqrstuvwxyz"
    encrypted = ''
    for char in message:
        if char in alphabet:
            pos = alphabet.index(char)
            encrypted = encrypted + cipher[pos]
        else:
            encrypted = encrypted + char
    return encrypted

def decrypt(encrypted, cipher):
    alphabet = "abcdefghijklmnopqrstuvwxyz"
    decrypted = ''
    for char in encrypted:
        if char in alphabet:
            pos = cipher.index(char)
            decrypted = decrypted + alphabet[pos]
        else:
            decrypted = decrypted + char
    return decrypted


cipher = "badcfehgjilknmporqtsvuxwzy"

encrypted = encrypt('hello world', cipher)
print(encrypted)

decrypted = decrypt(encrypted, cipher)
print(decrypted)
        </input>
            </program>
          </solution>
        </exercise>

        <exercise label="str_remove_dups">
          <statement>
            <p>Write a function called <c>remove_dups</c> that takes a string and creates a new string by only adding those characters that are not already present.  In other words, there will never be a duplicate letter added to the new string.</p>
          </statement>
          <program xml:id="str_remove_dups_editor" interactive="activecode" language="python">
            <input>
def remove_dups(a_string):
    """Returns unique characters in a_string."""
    # your code here


print(remove_dups("mississippi"))   #should print misp
</input>
<tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(remove_dups("pooh"),"poh","Tested remove_dups on string 'pooh'")
      self.assertEqual(remove_dups("mississippi"),"misp","Tested remove_dups on string 'mississippi'")
      self.assertEqual(remove_dups("potato"),"pota","Tested remove_dups on string 'potato'")
      self.assertEqual(remove_dups("bookkeeper"),"bokepr","Tested remove_dups on string 'bookkeeper'")
      self.assertEqual(remove_dups("oo"),"o","Tested remove_dups on string 'oo'")

myTests().main()
        </tests>
          </program>
        </exercise>

        <!--
        <exercise label="str_8_20">
          <statement>
            <p>Write a function called <c>rot13</c> that uses the Caesar cipher to encrypt a message.
                            The Caesar cipher works like a substitution cipher but each character is replaced
                            by the character 13 characters to &#x2018;its right' in the alphabet.  So for example
                            the letter a becomes the letter n.  If a letter is past the middle of the alphabet
                            then the counting wraps around to the letter a again, so n becomes a, o becomes b
                            and so on.  <em>Hint:</em> Whenever you talk about things wrapping around its a good idea
                            to think of modulo arithmetic.</p>
          </statement>
          <program interactive="activecode" language="python" xml:id="str_8_20_editor">
            <input>
def rot13(mess):
    # Your code here

print(rot13('abcde'))
print(rot13('nopqr'))
print(rot13(rot13('Since rot13 is symmetric you should see this message')))
        </input>
          </program>
          <solution>
            <program xml:id="str_q21_answer" language="python">
              <input>
def rot13(mess):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted = ''
    for char in mess:
        if char == ' ':
            encrypted = encrypted + ' '
        else:
            rotated_index = alphabet.index(char) + 13
            if rotated_index &lt; 26:
                encrypted = encrypted + alphabet[rotated_index]
            else:
                encrypted = encrypted + alphabet[rotated_index % 26]
    return encrypted

print(rot13('abcde'))
print(rot13('nopqr'))
print(rot13(rot13('since rot thirteen is symmetric you should see this message')))
        </input>
            </program>
          </solution>
        </exercise>
        -->
        <exercise label="str_format_money">
          <statement>
            <p>Modify this code so it prints each subtotal, the total cost, and average price to exactly two decimal places. Hint: use f-strings.</p>
          </statement>
          <program xml:id="str_format_money_editor" interactive="activecode" language="python">
            <input>
def checkout():
    """Print running subtotal, total, and average price of purchase.
    Program asks for input and accumulates subtotal until user
    enters zero for price.
    """
    total = 0
    count = 0
    moreItems = True
    while moreItems:
        price = float(input('Enter price of item (0 when done): '))
        if price != 0:
            count = count + 1
            total = total + price
            print('Subtotal: $', total)
        else:
            moreItems = False
    average = total / count
    print('Total items:', count)
    print('Total $', total)
    print('Average price per item: $', average)

checkout()
        </input>
          </program>
        </exercise>

</exercises>
