<?xml version="1.0"?>
<exercises xml:id="lists_exercises">
  <title>Exercises</title>

  <exercise label="ex_append_concat">
    <statement>
      <p>Create a list called <c>myList</c> with the following six items: 76, 92.3, <q>hello</q>, True, 4, 76. Begin with the empty list shown below, and add 6 statements to add each item, one per item. The first three statements should use the append method to append the item to the list, and the last three statements should use concatenation.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_2_editor">
      <code>

my_list = []

# Your code here
  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(my_list,[76, 92.3, 'hello', True, 4, 76],"my_list should contain the specified items")
        self.assertIn(".append(", self.getEditorText(), 'append method must be used')

myTests().main()
        </tests>
    </program>
    <solution>
      <program xml:id="ex_9_2_answer" language="python">
        <code>
my_list = []
my_list.append(76)
my_list.append(92.3)
my_list.append("hello")
my_list = my_list + [True]
my_list = my_list + [4]
my_list = my_list + [76]
print(my_list)
        </code>
      </program>
    </solution>
  </exercise>
  
  <exercise label="ex_list_methods">
    <statement>
      <p>Starting with the list of the previous exercise, write Python statements to do the following:</p>
      <p>
        <ol marker="a">
          <li>
            <p>Append <q>apple</q> and 76 to the list.</p>
          </li>
          <li>
            <p>Insert the value <q>cat</q> at position 3.</p>
          </li>
          <li>
            <p>Insert the value 99 at the start of the list.</p>
          </li>
          <li>
            <p>Find the index of <q>hello</q>.</p>
          </li>
          <li>
            <p>Count the number of 76s in the list.</p>
          </li>
          <li>
            <p>Remove the first occurrence of 76 from the list.</p>
          </li>
          <li>
            <p>Remove <c>True</c> from the list using <c>pop</c> with an <c>index</c> argument.</p>
          </li>
        </ol>
      </p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_list_methods_editor">
      <code>
my_list = [76, 92.3, 'hello', True, 4, 76]

# Your code here
        </code>
    </program>
    <solution>
      <program xml:id="ex_9_3_answer" language="python">
        <code>
my_list = [76, 92.3, 'hello', True, 4, 76]

my_list.append("apple")         # a
my_list.append(76)              # a
my_list.insert(3, "cat")        # b
my_list.insert(0, 99)           # c

print(my_list.index("hello"))   # d
print(my_list.count(76))        # e
my_list.remove(76)              # f
my_list.pop(my_list.index(True)) # g

print (my_list)
        </code>
      </program>
    </solution>
  </exercise>
  
  <exercise label="ex_average_list">
    <statement>
      <p>Write a function called <c>average</c> that takes a list of numbers as a parameter and returns the average of the numbers.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_list_avg_editor">
      <code>
def average(num_list):
    # Complete the function definition

  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(int(average([1, 3, 5, 7])), 4, "average([1, 3, 5, 7]) should be 4")
        self.assertEqual(int(average([9, 5, 4])), 6, "average([9, 5, 4]) should be 6")

myTests().main()
        </tests>
    </program>
    <solution>
      <program xml:id="ex_list_avg_solution" language="python">
        <code>
def average(num_list):

    total = 0
    for num in num_list:
        total = total + num

    return total / len(num_list)
        </code>
      </program>
    </solution>
  </exercise>
  
  <exercise label="ex_list_max">
    <statement>
      <p>Write a Python function named <c>max</c> that takes a parameter containing a nonempty list of integers and returns the maximum value.  (Note: there is a built-in function named <c>max</c>, but do not use it.)</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_list_max_editor">
      <code>
def max(num_list):
    # Complete the function

  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(max([3, 31, 5, 7]), 31, "max([3, 31, 5, 7]) should be 31")
        self.assertEqual(max([3, 13, 51, 7]), 51, "max([3, 13, 51, 7]) should be 51")

myTests().main()
        </tests>
    </program>
    <solution>
      <program xml:id="list_max_answer" language="python">
        <code>
def max(number_list):
    max_value = number_list[0]
    for element in number_list:
        if element &gt; max_value:
            max_value = element
    return max_value
        </code>
      </program>
    </solution>
  </exercise>
  
  <exercise label="ex_list_sum_squares">
    <statement>
      <p>Write a function <c>sum_of_squares(num_list)</c> that computes the sum of the squares of the numbers in the list <c>num_list</c>.  For example, <c>sum_of_squares([2, 3, 4])</c> should return 4&#160;+&#160;9&#160;+&#160;16, which is 29:</p>
    </statement>
    <program xml:id="ex_list_sum_squares_editor" interactive="activecode" language="python">
      <code>
def sum_of_squares(num_list):
    # your code here

  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_of_squares([2, 3, 4]), 29, "Tested sum_of_squares on input [2, 3, 4]")
        self.assertEqual(sum_of_squares([0, 1, -1]), 2, "Tested sum_of_squares on input [0, 1, -1]")
        self.assertEqual(sum_of_squares([5, 12, 14]), 365, "Tested sum_of_squares on input [5, 12, 14]")

myTests().main()
        </tests>
    </program>
  </exercise>
  
  <exercise label="ex_list_count_odd">
    <statement>
      <p>Write a function to count how many odd numbers are in a list.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_list_count_odd_editor">
      <code>
def count_odd(num_list):
    # your code here

  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(count_odd([1, 3, 5, 7, 9]), 5, "Tested count_odd on input [1, 3, 5, 7, 9]")
        self.assertEqual(count_odd([1, 2, 3, 4, 5]), 3, "Tested count_odd on input [-1, -2, -3, -4, -5]")
        self.assertEqual(count_odd([2, 4, 6, 8, 10]), 0, "Tested count_odd on input [2, 4, 6, 8, 10]")
        self.assertEqual(count_odd([0, -1, 12, -33]), 2, "Tested count_odd on input [0, -1, 12, -33]")

myTests().main()
        </tests>
    </program>
    <solution>
      <program xml:id="ex_list_count_odd_answer" language="python">
        <code>
import random

def count_odd(num_list):
    odd = 0
    for element in num_list:
        if element % 2 != 0:
            odd = odd + 1
    return odd

# make a random list to test the function
numbers = []
for i in range(10):
    numbers.append(random.randint(0, 1000))

print(numbers) # show the random numbers
print(count_odd(numbers))
        </code>
      </program>
    </solution>
  </exercise>
  
  <exercise label="ex_list_sum_evens">
    <statement>
      <p>Sum up all the even numbers in a list.</p>
    </statement>
    <program xml:id="ex_list_sum_evens_editor" interactive="activecode" language="python">
      <code>
def sum_even(number_list):
    # your code here

  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_even([1, 3, 5, 7, 9]), 0, "Tested sum_even on input [1, 3, 5, 7, 9]")
        self.assertEqual(sum_even([-1, -2, -3, -4, -5]), -6, "Tested sum_even on input [-1, -2, -3, -4, -5]")
        self.assertEqual(sum_even([2, 4, 6, 7, 9]), 12, "Tested sum_even on input [2, 4, 6, 7, 9]")
        self.assertEqual(sum_even([0, 1, 12, 33]), 12, "Tested sum_even on input [0, 1, 12, 33]")

myTests().main()
        </tests>
    </program>
  </exercise>
  
  <exercise label="ex_list_product_negatives">
    <statement>
      <p>Calculate the product of all the negative numbers in a list. For example, <c>product_negatives([-3, 2, -4, 7])</c> should be 12. If a list contains no negative numbers, return 1.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_list_product_negatives_editor">
      <code>
def product_negatives(number_list):
    # your code here

  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(product_negatives([-1, -2, -3, -4, -5]), -120, "Tested product_negatives on input [-1, -2, -3, -4, -5]")
        self.assertEqual(product_negatives([1, -3, 5, -7, 9]), 21, "Tested product_negatives on input [1, -3, 5, -7, 9]")
        self.assertEqual(product_negatives([-2, -4, 6, -7, 9]), -56, "Tested product_negatives on input [-2, -4, 6, -7, 9]")
        self.assertEqual(product_negatives([1, 2, 3, 4]), 1, "Tested product_negatives on input [1, 2, 3, 4]")

myTests().main()
        </tests>
    </program>
    <solution>
      <program xml:id="ex_list_product_negatives_answer" language="python">
        <code>
import random

def product_negatives(number_list):
    product = 1
    for element in number_list:
        if element &lt; 0:
            product = product * element
    return product

numbers = []
for i in range(10):
    numbers.append(random.randrange(-10, 10))

print(numbers) # show numbers
print(product_negatives(numbers))
        </code>
      </program>
    </solution>
  </exercise>
  
  <exercise label="ex_list_word_length">
    <statement>
      <p>Count how many words in a list have length 5.</p>
    </statement>
    <program xml:id="ex_list_word_length_editor" interactive="activecode" language="python">
      <code>
def count_words(word_list):
    # your code here
    
print(count_words(["plank", "wood", "board", "nail"]) # 2
print(count_words(["cat", "mouse", "tapir", "puppy", "dog"]) # 3
        </code>
    </program>
  </exercise>
  
  <exercise label="ex_list_implement_methods">
    <statement>
      <p>Although Python provides us with many list methods, it is good practice and very instructive to think about how they are implemented.  Implement Python functions that work like the following:</p>
      <p>
        <ol marker="a">
          <li>
            <p><c>count</c></p>
          </li>
          <li>
            <p><c>in</c> (you will have to name this function <c>is_in</c> because <c>in</c> is a reserved word)</p>
          </li>
          <li>
            <p><c>reverse</c></p>
          </li>
          <li>
            <p><c>index</c></p>
          </li>
          <li>
            <p><c>insert</c></p>
          </li>
        </ol>
      </p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_list_implement_methods_editor">
      <code>
# Your functions go here..

a_list = [0, 1, 1, 2, 2, 3, 4, 5, 6, 7, 8, 9]
print(count(1, a_list))
print(is_in(4, a_list))
print(reverse(a_list))
print(index(2, a_list))
print(insert('cat', 4, a_list))
</code>
    </program>
    <solution>
      <program xml:id="ex_list_implement_methods_answer" language="python">
        <code>
def count(obj, a_list):
    count = 0
    for element in a_list:
        if element == obj:
            count = count + 1
    return count

def is_in(obj, a_list):  # cannot be called in() because in is a reserved keyword
    for element in a_list:
        if element == obj:
            return True
    return False

def reverse(a_list):
    reversed = []
    for i in range(len(a_list) - 1, -1, -1): # step through the original list backwards
        reversed.append(a_list[i])
    return reversed

def index(obj, a_list):
    for i in range(len(a_list)):
        if a_list[i] == obj:
            return i
    return -1

def insert(obj, index, a_list):
    new_list = []
    for i in range(len(a_list)):
        if i == index:
            new_list.append(obj)
        new_list.append(a_list[i])
    return new_list

a_list = [0, 1, 1, 2, 2, 3, 4, 5, 6, 7, 8, 9]
print(count(1, a_list))
print(is_in(4, a_list))
print(reverse(a_list))
print(index(2, a_list))
print(insert('cat', 4, a_list))
        </code>
      </program>
    </solution>
  </exercise>
  
  <!--
  <exercise label="ex_list_replace">
    <statement>
      <p>Write a function <c>replace(s, old, new)</c> that replaces all occurences of
                <c>old</c> with <c>new</c> in a string <c>s</c>:</p>
      <pre>test(replace('Mississippi', 'i', 'I'), 'MIssIssIppI')

s = 'I love spom!  Spom is my favorite food.  Spom, spom, spom, yum!'
test(replace(s, 'om', 'am'),
       'I love spam!  Spam is my favorite food.  Spam, spam, spam, yum!')

test(replace(s, 'o', 'a'),
       'I lave spam!  Spam is my favarite faad.  Spam, spam, spam, yum!')</pre>
      <p><em>Hint</em>: use the <c>split</c> and <c>join</c> methods.</p>
    </statement>
    <program xml:id="ex_9_13_editor" interactive="activecode" language="python">
      <code>
def replace(s, old, new):
    # your code here

  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(replace('Mississippi','i','I'),'MIssIssIppI',"Tested replace on input 'Mississippi','i','I'")
        self.assertEqual(replace('Bookkeeper','e','A'),'BookkAApAr',"Tested failed on input 'Bookkeeper','e','A'")
        self.assertEqual(replace('Deeded','e','q'),'Dqqdqd',"Tested failed on input 'Deeded','e','q'")

myTests().main()
        </tests>
    </program>
  </exercise>
  -->

</exercises>
