<?xml version="1.0"?>
<section xml:id="lists_pure-functions">
  <title>Pure Functions</title>
  
  <p>The revised version of the <c>remove_negative</c> function that we developed in <xref ref="lists_deletion-redux"/> is a <term>pure function</term>. A pure function does not produce <term>side effects</term> such as printing to the screen or changing its parameters. It communicates with the calling program only through parameters (which it does not modify) and a return value.
  </p>
  
  <exercise xml:id="lists_pure_function_ex">
    <statement>
      <p>
        A word is <q>vowelicious</q> if it contains four or more vowels. (If you never heard this word before, it<rsq/>s because I made it up just now for this book.) For example, <q>mountain</q> and <q>bookbinder</q> are vowelicious; <q>trailer</q> and <q>reading</q> are not.
      </p>
      
      <p>Write a program that takes a list of words and creates a new list containing only the vowelicious words. Your program should have two pure functions:</p>
      
      <p>
        <ul>
          <li><p><c>count_vowels</c>, which takes a string and returns the number of vowels it contains</p></li>
          <li><p><c>keep_vowelicious</c>, which takes a list of words and returns a new list containing only the words that have four or more vowels.</p></li>
        </ul>
      </p>
    </statement>
  <program xml:id="vowelicious_program" interactive="activecode" language="python">
    <input>
def count_vowels(a_string):
    """ Return the number of vowels in the parameter string.
    Vowels are a, e, i, o, and u (upper and lower case).
    The letter y is not considered to be a vowel in this program.
    """
    # your code here...
    
def keep_vowelicious(a_list):
    """Keep only words with four or more vowels.
    Given a list of words, return a new list that contains
    only words that have four or more vowels.
    """
    # your code here...

words = ["beautiful", "locomotive", "friendly", "automation", "treadmill", "kangaroo"]
kept_words = keep_vowelicious(words)
print(f"The vowelicious words are {kept_words}")
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(kept_words, ['beautiful', 'locomotive', 'automation', 'kangaroo'], "kept_words should contain the specified items")
        self.assertEqual(words, ["beautiful", "locomotive", "friendly", "automation", "treadmill", "kangaroo"], "Original word list must remain unchanged.")
myTests().main()
    </input>
  </program>
  <hint>
  <p>The <c>keep_vowelicious</c> function should call the <c>count_vowels</c> function.</p>
</hint>

</exercise>
</section>
