<?xml version="1.0"?>
<section xml:id="lists_the-accumulator-pattern-with-lists">
  <title>The Accumulator Pattern with Lists</title>
  <introduction>

    <p>
      On the preceding page, we added all the items in the list to the accumulator. Sometimes when we're accumulating, we don't want to add to our accumulator every time we iterate. Consider, for example, the following program which counts the number of names with more than 3 letters.
    </p>
  <program xml:id="accum_count" interactive="activecode" language="python">
    <code>
long_names = 0
for name in ["Joe", "Sally", "Amy", "Brad"]:
    if len(name) &gt; 3:
        long_names += 1
print(long_names)
        </code>
  </program>
  <p>
    Here, we <term>initialize</term> the accumulator variable to be zero on line 1.
  </p>
  
  <p>
    We <term>iterate</term> through the sequence (line 2).
  </p>
  
  <p>
    The <term>update</term> step happens in two parts. First, we check to see if the name is longer than 3 letters. If so, then we increment the accumulator variable <c>long_names</c> (on line 4) by adding one to it.
  </p>
  
  <p>
    At the end, we have accumulated the total number of long names.
  </p>
  
  <p>
    We can use conditionals to also count if particular items are in a string or list. The following code finds all occurrences of vowels in a string.
  </p>
  
  <program xml:id="accum_vowels" interactive="activecode" language="python">
    <code>
s = "what if we went to the zoo"
num_vowels = 0
for char in s:
    if char in 'aeiou':
        num_vowels += 1
print(num_vowels)
        </code>
  </program>
  
</introduction>
  <subsection xml:id="lists_accumulating-the-max-value">
    <title>Accumulating the Maximum Value</title>
    <p>
      We can also use the accumulation pattern with conditionals to find the maximum or minimum value. Instead of continuing to build up the accumulator value like we have when counting or finding a sum, we can reassign the accumulator variable to a different value.
  </p>
  
    <p>
      The following example shows how we can get the maximum value from a list of integers.
  </p>
  
    <program xml:id="accum_max_value" interactive="activecode" language="python">
      <code>
nums = [9, 3, 8, 11, 5, 29, 2]
biggest_num = nums[0]
for n in nums:
    if n &gt; biggest_num:
        biggest_num = n
print(biggest_num)
        </code>
    </program>
    <p>
      Here, we initialize <c>biggest_num</c> to the first element in our list. Why didn't we set it to zero? Because we cannot assume that there will never be any negative numbers in the list. By assigning <c>nums[0]</c> to <c>biggest_num</c>, we are at least guaranteed that the result of our program will be one of the items in the list.
  </p>
    <p>
      In the <c>for</c> loop, we check to see if the current value of <c>n</c> is greater than the current value of <c>biggest_num</c>. If it is, then we want to <term>update</term> <c>biggest_num</c> so that it now is assigned the higher number. Otherwise, we do nothing and continue the <c>for</c> loop.
  </p>
  </subsection>
  
  <subsection xml:id="lists_accumulating-a-string-result">
    <title>Accumulating a String Result</title>
    <p>
      The accumulator pattern can be used to convert a list of items to a string.
    </p>
    <p>
      Consider the following program:
    </p>
    
    <program xml:id="accum_string" interactive="activecode" language="python">
      <code>
scores = [85, 95, 70]
result = ''
for score in scores:
   result = result + str(score) + ','

print("The scores are " + result)
        </code>
    </program>
    <p>
      Here, the accumulator variable is <c>result</c>. Each time through the loop, the program concatenates the current contents of <c>result</c> with the comma separator and a score from the list, and updates the <c>result</c> with the new value. Use CodeLens to step through this example to see it in action.
    </p>
    <p>
      The output of the program has some undesirable formatting problems: there is a trailing comma instead of a period, and there are no spaces between the items. The next activity lets you work to correct those problems.
  </p>
  
    <exercise label="accum_better_string">
      <statement>
        <p> Let's work to improve the formatting of the sentence produced by the program above. Revise the following code so that it outputs the sentence:
        </p>
        <pre>The scores are 85, 95, and 70.</pre>
      </statement>
      <program interactive="activecode" language="python" xml:id="accum_better_string_editor">
        <code>
scores = [85, 95, 70]
result = ''
for score in scores:
   result = result + str(score) + ','

print("The scores are " + result)
  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(self.getOutput().rstrip('\n'), 'The scores are 85, 95, and 70.', "Output should be 'The scores are 85, 95, and 70.'")

myTests().main()
        </tests>
      </program>
      <hint>
        <p>
          Try changing the loop so that it does not process the final score. Handle that last score separately, after the loop finishes.
        </p>
      </hint>
      <solution>
          This solution works by iterating over all of the scores in the list except the last, and dealing with that one separately.

    <program xml:id="accum_better_string_solution" language="python"><code>
scores = [85, 95, 70]
result = ''
for score in scores[:-1]:
   result = result + str(score) + ', '

# Now, append the last score
result = result + 'and ' + str(scores[-1]) + '.'

print("The scores are " + result)
        </code></program>
                </solution>
    </exercise>

    <exercise label="accum_pattern_q1">
      <statement>
        <p>
          What is printed by the following statements?
        </p>
        <program language="python">
          <code>
s = "We are learning!"
x = 0
for char in s:
    if char in 'abcde':
        x = x + 1
print(x)
</code>
        </program>
      </statement>
      <choices>
        <choice>
          <statement>
            <p>2</p>
          </statement>
          <feedback>Though only two of the letters in the list are found, we count them each time they appear.</feedback>
        </choice>
        <choice correct="yes">
          <statement>
            <p>5</p>
          </statement>
          <feedback>Yes, we add to <c>x</c> each time we come across a letter in the list.</feedback>
        </choice>
        <choice>
          <statement>
            <p>0</p>
          </statement>
          <feedback>Check again what the conditional is evaluating. The value of <c>ch</c> will be a character in the string <c>s</c>, so what will happen in the <c>if</c> statement?</feedback>
        </choice>
        <choice>
          <statement>
            <p>There is an error in the code so it cannot run.</p>
          </statement>
          <feedback>There are no errors in this code.</feedback>
        </choice>
      </choices>
    </exercise>

    <exercise label="accum_pattern_q2">
      <statement>
        <p>
    What is printed by the following statements?
  </p>
        <program language="python">
          <code>
my_list= [5, 2, 1, 4, 9, 10]
min_value = 0
for item in my_list:
   if item &lt; min_value:
       min_value = item
print(min_value)
</code>
        </program>
      </statement>
      <choices>
        <choice>
          <statement>
            <p>10</p>
          </statement>
          <feedback>Not quite. What is the conditional checking?</feedback>
        </choice>
        <choice>
          <statement>
            <p>1</p>
          </statement>
          <feedback><c>min_value</c> was set to a number that was smaller than any of the numbers in the list, so it was never updated in the <c>for</c> loop.</feedback>
        </choice>
        <choice correct="yes">
          <statement>
            <p>0</p>
          </statement>
          <feedback>Yes, <c>min_value</c> was set to a number that was smaller than any of the numbers in the list, so it was never updated in the <c>for</c> loop.</feedback>
        </choice>
        <choice>
          <statement>
            <p>There is an error in the code so it cannot run.</p>
          </statement>
          <feedback>The code does not have an error that would prevent it from running.</feedback>
        </choice>
      </choices>
    </exercise>
    
    <exercise label="accum_verb_forms">
      <statement>
        <p><term>Challenge</term> For each word in <c>words</c>, add &#x2018;d' to the end of the word if the word ends in <q>e</q> to make it past tense. Otherwise, add &#x2018;ed' to make it past tense. Save these past tense words to a list called <c>past_tense</c>.</p>
      </statement>
      <program xml:id="accum_verb_forms" interactive="activecode" language="python">
        <code>
words = ["adopt", "bake", "beam", "confide", "grill", "plant", "time", "wave", "wish"]

  </code>
  <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testNine(self):
      self.assertEqual(past_tense, ['adopted', 'baked', 'beamed', 'confided', 'grilled', 'planted', 'timed', 'waved', 'wished'], "Testing that the past_tense list is correct.")
      self.assertIn("else", self.getEditorText(), "Testing output (Don't worry about actual and expected values).")
      self.assertIn("for", self.getEditorText(), "Testing output (Don't worry about actual and expected values).")

myTests().main()
        </tests>
      </program>
    </exercise>
    
    <p>Now that we have finished our in-depth review of the accumulator pattern, let's look at lists in detail, first describing the things they have in common with strings.</p>
 
  </subsection>
</section>
